// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Option - Definition

Description:
`Option` is a standard library enum used to represent the presence or absence of a value.
It has two variants: `Some(T)` and `None`.

Your task is to:
1. Define a variable `some_number` as `Option<u16>` containing the value 5.
2. Define a variable `absent_number` as `Option<u16>` containing nothing.
*/

fn main() {
    // TODO: Define `some_number`
    // let some_number: Option<u16> = ...;

    // TODO: Define `absent_number`
    // let absent_number: Option<u16> = ...;

    // println!("Number: {:?}, Absent: {:?}", some_number, absent_number);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
