// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Modules - Definition

Description:
Modules organize code into groups.
You can define a module using `mod name { ... }`.

Your task is to define a module named `sausage_factory` and a public function `make_sausage` inside it.
*/

// TODO: Define module and function
// mod sausage_factory { ... }

fn main() {
    // sausage_factory::make_sausage();
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
