// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Functions

Description:
Functions are the building blocks of Rust code.
The `main` function is the entry point of every program.

Your task is to define a new function named `call_me` and call it from `main`.

Hints:
1. `fn call_me() { ... }`
*/

fn main() {
    // TODO: Call the function `call_me`
}

// TODO: Define the function `call_me`
