// I AM NOT DONE

/*
Difficulty: ⭐
Topic: HashMap - Creation

Description:
A `HashMap` stores keys and values.
It requires `use std::collections::HashMap;`.

Your task is to:
1. Create a new HashMap called `scores`.
2. Insert ("Blue", 10) and ("Yellow", 50).
*/

use std::collections::HashMap;

fn main() {
    // TODO: Create a new HashMap
    // let mut scores = ...;

    // TODO: Insert values
    // scores.insert(String::from("Blue"), 10);
    // scores.insert(String::from("Yellow"), 50);

    // println!("{:?}", scores);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
