// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Vectors - Creation

Description:
Vectors are re-sizable arrays.
You can create them with `Vec::new()` or the `vec![]` macro.

Your task is to:
1. Create a vector `v1` using `Vec::new()`.
2. Create a vector `v2` containing [1, 2, 3] using `vec![]`.
*/

fn main() {
    // TODO: Create v1
    // let v1: Vec<i32> = ...;

    // TODO: Create v2
    // let v2 = ...;

    // println!("v1: {:?}, v2: {:?}", v1, v2);
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
