// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Match

Description:
The `match` control flow construct is like a switch statement in other languages, but more powerful.
It allows you to compare a value against a series of patterns and execute code based on which pattern matches.

Your task is to write a match expression that checks `number`:
- If 1, print "One"
- If 2, print "Two"
- If 3, print "Three"
- Use `_` (underscore) for any other value and print "Other".
*/

fn main() {
    let number = 2;

    // TODO: Write the match expression
    // match number { ... }
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
