// I AM NOT DONE

/*
Difficulty: â­
Topic: Macros - Definition

Description:
Macros allow you to write code that writes other code.
Declarative macros are defined using `macro_rules!`.

Your task is to define a simple macro `my_macro` that prints "Hello, Macro!".
*/

// TODO: Define my_macro
// macro_rules! my_macro {
//     () => {
//         println!("Hello, Macro!");
//     };
// }

fn main() {
    // my_macro!();
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
