// I AM NOT DONE

/*
Difficulty: ⭐
Topic: Traits - Definition

Description:
Traits define shared behavior. They are similar to interfaces in other languages.

Your task is to define a trait `Animal` with a method `sound` that returns a `String`.
*/

// TODO: Define the trait Animal
// trait Animal { ... }

fn main() {
    // Just to check compilation (cannot instantiate trait directly)
}

#[cfg(test)]
mod tests {
    #[test]
    fn test_main_runs() {
        super::main();
    }
}
